## Python

IPython magic are special commands that can be used to interact with your System.

Enter the following into an IPython session:

```python
(an_integer, 
 a_list,
 a_dictionary,
 a_set,
 a_tuple) = 1, [1,2,3], {'k': 1}, {1,2,2,3}, (1,2)
```

### Writing a Function

Functions in python are written using the keyword `def`.

    def function_name(arg1, arg2):
        output_1 = do_something_with(arg1)
        output_2 = do_something_with(arg2)
        return 

```python
from sys import getsizeof

def sizeof_and_value(variable):
    return (getsizeof(variable), variable)
```

```python
sizeof_and_value(an_integer)
```

Output:
```
(28, 1)
```

```python
sizeof_and_value(a_tuple)
```

Output:
```
(64, (1, 2))
```

#### Write a Function

Write a function named `type_and_value` that returns the value and the type of a variable that is passed to it.

Use the block below to define the function. 

<!-- grade_id: cell-c48f82b1aae3c2f5, points: 0, grade: false, solution: true -->
```python
def type_and_value(var):
    """return the type and value of a variable.
    """
    ### BEGIN SOLUTION
    return type(var), var
    ### END SOLUTION
```

Make sure that your function can pass this test:

<!-- grade_id: cell-80a7fd4b5e0528b5, points: 2, grade: true, solution: false -->
```python
assert type_and_value(an_integer) == (int, 1)

### BEGIN HIDDEN TESTS
assert type_and_value(an_integer) == (int, 1)
assert type_and_value(a_list) == (list,[1,2,3])
assert type_and_value(a_dictionary) == (dict,{'k': 1})
assert type_and_value(a_set) == (set,{1,2,2,3})
assert type_and_value(a_tuple) == (tuple,(1,2))
### END HIDDEN TESTS
```

### IPython Magic

### `%whos`

Prints a table with some basic details about each identifier you have defined interactively.


```python
%whos 
```

Output:
```
Variable           Type                          Data/Info
----------------------------------------------------------
a_dictionary       dict                          n=1
a_list             list                          n=3
a_set              set                           {1, 2, 3}
a_tuple            tuple                         n=2
an_integer         int                           1
getsizeof          builtin_function_or_method    <built-in function getsizeof>
sizeof_and_value   function                      <function sizeof_and_value at 0x7f18d0307e18>
type_and_value     function                      <function type_and_value at 0x7f18d0307ea0>
```

### Bash in IPython

Bash is a command line language used to interactive with an operating system.

Some simple Bash commands can be run in IPython/Jupyter, including

- `pwd`
- `cd`
- `ls`

#### `pwd` - print working directory

```python
%pwd
```

Output:
```
'/home/jovyan/introductiontodatascience/source/01-seeds'
```

#### `cd` - change directory

```python
%cd src/
```

Output:
```
[Errno 2] No such file or directory: 'src/'
/home/jovyan/introductiontodatascience/source/01-seeds
```

```python
%pwd
```

Output:
```
'/home/jovyan/introductiontodatascience/source/01-seeds'
```

#### `ls` - list files

```python
%ls
```

Output:
```
01-00-seeds.ipynb  01-01-interactive-programming.ipynb  01-03.ipynb
01-00-seeds.rst    01-02-Python.ipynb                   01-04.ipynb
```

<!-- grade_id: cell-6a5baeb91a451736, points: 2, grade: true, solution: false -->
```python
# HIDDEN TEST

### BEGIN HIDDEN TESTS
import os
assert os.getcwd().split('/')[-1] == 'src'
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-12-1f0ebf75cbe5> in <module>()
      3 ### BEGIN HIDDEN TESTS
      4 import os
----> 5 assert os.getcwd().split('/')[-1] == 'src'
      6 ### END HIDDEN TESTS

AssertionError:
```

### IPython Magic commands

There are many IPython magic commands, but some of the more useful are

- `run`
- `matplotlib inline`
- `whos`


```python
%run a_simple_script.py
```

```python
%matplotlib inline
```

```python
%run a_simple_script.py
```

#### Why didn't the image show up the first time?